# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

```python
import pandas as pd

# We are gonna read the dataset and see if there's duplicated in the id to prove that are uniques. Like we read in the 
# documentation the field 'SEQNUMC' contain the information about the mother's education of each children.

df = pd.read_csv('assets/NISPUF17.csv')
```

```python
df['EDUC1'].value_counts()
```

Output:
```
4    13656
3     6999
2     4906
1     2904
Name: EDUC1, dtype: int64
```

```python
total = df['EDUC1'].count()

less =    df['EDUC1'].value_counts()[1] / total
high_s =  df['EDUC1'].value_counts()[2] / total
m_high =  df['EDUC1'].value_counts()[3] / total
college = df['EDUC1'].value_counts()[4] / total


result = {"less than high school": less,
    "high school": high_s,
    "more than high school but not college": m_high,
    "college": college}

print(result)
```

Output:
```
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    
    total = df['EDUC1'].count()

    less =    df['EDUC1'].value_counts()[1] / total
    high_s =  df['EDUC1'].value_counts()[2] / total
    m_high =  df['EDUC1'].value_counts()[3] / total
    college = df['EDUC1'].value_counts()[4] / total


    result = {"less than high school": less,
        "high school": high_s,
        "more than high school but not college": m_high,
        "college": college}
    
    
    return result
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
# Explore the field related to the breastmilk of each children and see the groups

# Knowing that the defined groups are 1 'Yes' and 2 'No', we'll make a mask and get the mean() of the defined variable. After, look
# for the field I know that which indicates the number of vaccines of influenza givens to the child is 'P_NUMFLU'.


y_milk = df.loc[df['CBF_01'] == 1, 'P_NUMFLU'].mean()
no_milk = df.loc[df['CBF_01'] == 2, 'P_NUMFLU'].mean()

print(y_milk, no_milk)
```

Output:
```
1.8799187420058687 1.5963945918878317
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():

    y_milk = df.loc[df['CBF_01'] == 1, 'P_NUMFLU'].mean()
    no_milk = df.loc[df['CBF_01'] == 2, 'P_NUMFLU'].mean()
    
    tuple = (y_milk, no_milk)
    return (y_milk, no_milk)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
# We make a Dataframe with the total of man and womans who has or not chickenpox to look for the relationship later.

total_m = df.loc[(df['P_NUMVRC'] >= 1) & (df['SEX'] == 1) & (df['HAD_CPOX'] < 3)]

total_w = df.loc[(df['P_NUMVRC'] >= 1) & (df['SEX'] == 2) & (df['HAD_CPOX'] < 3)]

# Once we have the totals we can get the ratio of these relationship about their own groups

man = len(total_m[total_m['HAD_CPOX'] == 1]) / len(total_m[total_m['HAD_CPOX'] == 2])

man = len(total_m[total_m['HAD_CPOX'] == 1]) / len(total_m[total_m['HAD_CPOX'] == 2])

# Ratio of womans who contracted chickenpox even when we were vaccinated against it about the total of womans vaccinated

woman = len(total_w[total_w['HAD_CPOX'] == 1]) / len(total_w[total_w['HAD_CPOX'] == 2])



len(total_w[total_w['HAD_CPOX'] == 1]) / len(total_w[total_w['HAD_CPOX'] == 2])

```

Output:
```
0.0077918259335489565
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    
    total_m = df.loc[(df['P_NUMVRC'] >= 1) & (df['SEX'] == 1) & (df['HAD_CPOX'] < 3)]

    total_w = df.loc[(df['P_NUMVRC'] >= 1) & (df['SEX'] == 2) & (df['HAD_CPOX'] < 3)]

    man = len(total_m[total_m['HAD_CPOX'] == 1]) / len(total_m[total_m['HAD_CPOX'] == 2])

    woman = len(total_w[total_w['HAD_CPOX'] == 1]) / len(total_w[total_w['HAD_CPOX'] == 2])
    
    result = {"male":   man,
              "female": woman}
    return result
```

```python
chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
import scipy.stats as stats
import numpy as np

n_df = df.loc[ df['HAD_CPOX'] < 3 , ['P_NUMVRC', 'HAD_CPOX']]

n_df.dropna(subset= ['P_NUMVRC', 'HAD_CPOX'], inplace = True)

n_df["P_NUMVRC"] = n_df["P_NUMVRC"].astype(int)

corr, pval=stats.pearsonr(n_df["HAD_CPOX"],n_df["P_NUMVRC"])

print(corr, pval)
```

Output:
```
0.07044873460147986 2.7780263182916748e-18
```

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
# Make a new dataset with the fields of interest

    n_df = df.loc[ df['HAD_CPOX'] < 3 , ['P_NUMVRC', 'HAD_CPOX']]

# Clean NaN and change the type to avoid errors when we make the pearson correlation

    n_df.dropna(subset= ['P_NUMVRC', 'HAD_CPOX'], inplace = True)
    n_df["P_NUMVRC"] = n_df["P_NUMVRC"].astype(int)

    corr, pval=stats.pearsonr(n_df["HAD_CPOX"],n_df["P_NUMVRC"])

    return corr
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```